## Codio Activity 1.2: Exploring a DataFrame 

**Expected Time: 30 Minutes**

**Total Points: 12 Points**

In this exercise you will use basic summarization methods on a pandas DataFrame.

## Index:

- [Problem 1](#Problem-1)
- [Problem 2](#Problem-2)
- [Problem 3](#Problem-3)

```python
import pandas as pd
import numpy as np
```

```python
nba = pd.read_csv('/Users/jaikumar/Desktop/Berkeley_ML/codio_1.2_solution/data/nba.csv', index_col = 0)
```

```python
nba.isnull().sum()
```

Output:
```
Name         0
Team         0
Position     0
Height       0
College     84
Number       0
Age          0
Weight       0
Salary       0
dtype: int64
```

```python
nba[:4]
```

Output:
```
Name            Team Position Height            College  Number  \
0  Avery Bradley  Boston Celtics       PG    6-2              Texas     0.0   
1    Jae Crowder  Boston Celtics       SF    6-6          Marquette    99.0   
2   John Holland  Boston Celtics       SG    6-5  Boston University    30.0   
3    R.J. Hunter  Boston Celtics       SG    6-5      Georgia State    28.0   

    Age  Weight     Salary  
0  25.0   180.0  7730337.0  
1  25.0   235.0  6796117.0  
2  27.0   205.0        0.0  
3  22.0   185.0  1148640.0
```

```python
nba[nba['Weight']>225]
```

Output:
```
Name            Team Position Height    College  Number   Age  \
1       Jae Crowder  Boston Celtics       SF    6-6  Marquette    99.0  25.0   
4     Jonas Jerebko  Boston Celtics       PF   6-10        NaN     8.0  29.0   
5      Amir Johnson  Boston Celtics       PF    6-9        NaN    90.0  29.0   
6     Jordan Mickey  Boston Celtics       PF    6-8        LSU    55.0  21.0   
7      Kelly Olynyk  Boston Celtics        C    7-0    Gonzaga    41.0  25.0   
..              ...             ...      ...    ...        ...     ...   ...   
448  Gordon Hayward       Utah Jazz       SF    6-8     Butler    20.0  26.0   
450      Joe Ingles       Utah Jazz       SF    6-8        NaN     2.0  28.0   
452      Trey Lyles       Utah Jazz       PF   6-10   Kentucky    41.0  20.0   
455    Tibor Pleiss       Utah Jazz        C    7-3        NaN    21.0  26.0   
456     Jeff Withey       Utah Jazz        C    7-0     Kansas    24.0  26.0   

     Weight      Salary  
1     235.0   6796117.0  
4 
```

```python
nba.loc[:3]
```

Output:
```
Name            Team Position Height            College  Number  \
0  Avery Bradley  Boston Celtics       PG    6-2              Texas     0.0   
1    Jae Crowder  Boston Celtics       SF    6-6          Marquette    99.0   
2   John Holland  Boston Celtics       SG    6-5  Boston University    30.0   
3    R.J. Hunter  Boston Celtics       SG    6-5      Georgia State    28.0   

    Age  Weight     Salary  
0  25.0   180.0  7730337.0  
1  25.0   235.0  6796117.0  
2  27.0   205.0        0.0  
3  22.0   185.0  1148640.0
```

```python
nba.loc[nba['Weight']>280,['Name', 'College']]
nba.loc[[5,6,9]]
```

Output:
```
Name            Team Position Height         College  Number  \
5   Amir Johnson  Boston Celtics       PF    6-9             NaN    90.0   
6  Jordan Mickey  Boston Celtics       PF    6-8             LSU    55.0   
9   Marcus Smart  Boston Celtics       PG    6-4  Oklahoma State    36.0   

    Age  Weight      Salary  
5  29.0   240.0  12000000.0  
6  21.0   235.0   1170960.0  
9  22.0   220.0   3431040.0
```

```python
nba.iloc[:3, :3]
```

Output:
```
Name            Team Position
0  Avery Bradley  Boston Celtics       PG
1    Jae Crowder  Boston Celtics       SF
2   John Holland  Boston Celtics       SG
```

```python
nba.iloc[[0,3,10], :3]
```

Output:
```
Name            Team Position
0     Avery Bradley  Boston Celtics       PG
3       R.J. Hunter  Boston Celtics       SG
10  Jared Sullinger  Boston Celtics        C
```

[Back to top](#Index:) 

## Problem 1

### Describing the data

**4 Points**

Using the DataFrame `nba` loaded above, get a summary of the numeric variables using the `.describe` method.  Assign your results as a DataFrame to `nba_summary` below. 

<!-- grade_id: cell-cb8365b624aba7a4, points: 0, grade: false, solution: true -->
```python
### GRADED

nba_summary = None

### BEGIN SOLUTION
nba_summary = nba.describe()
type(nba_summary)
### END SOLUTION

# Answer check
print(nba_summary)
print(type(nba_summary))
```

Output:
```
Number         Age      Weight        Salary
count  457.000000  457.000000  457.000000  4.570000e+02
mean    17.678337   26.938731  221.522976  4.726121e+06
std     15.966090    4.404016   26.368343  5.218945e+06
min      0.000000   19.000000  161.000000  0.000000e+00
25%      5.000000   24.000000  200.000000  1.000000e+06
50%     13.000000   26.000000  220.000000  2.658240e+06
75%     25.000000   30.000000  240.000000  6.331404e+06
max     99.000000   40.000000  307.000000  2.500000e+07
<class 'pandas.core.frame.DataFrame'>
```

<!-- grade_id: cell-ff1d43a9f019258d, points: 4, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
nba_summary_ = nba.describe()
#
#
#
assert type(nba_summary_) == type(nba_summary)
pd.testing.assert_frame_equal(nba_summary, nba_summary_)
### END HIDDEN TESTS
```

[Back to top](#Index:) 

### Problem 2

### Getting the missing data information 

**4 Points**

In the National Basketball Association (NBA), there are some athletes that do not take the usual career path of playing for a college and then become a professional athlete. Some athletes skip college and go professional directly out of high school. How many athletes in this dataframe skipped college and went straight to the professional level?

To find this answer, create a series with the counts of null values for each column in the `nba` DataFrame.  Assign your solution as a pandas Series to `nba_nulls` below.  **HINT**: Use the `.isnull` and `.sum` methods together. 

<!-- grade_id: cell-a78e6c5da54a7081, points: 0, grade: false, solution: true -->
```python
### GRADED

nba_nulls = None

### BEGIN SOLUTION
nba_nulls = nba.isnull().sum()
type(nba_nulls)
### END SOLUTION

# Answer check
print(nba_nulls)
print(type(nba_nulls))
```

Output:
```
Name         0
Team         0
Position     0
Height       0
College     84
Number       0
Age          0
Weight       0
Salary       0
dtype: int64
<class 'pandas.core.series.Series'>
```

<!-- grade_id: cell-8f2d38bd42b42677, points: 4, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
nba_nulls_ = nba.isnull().sum()
#
#
#
assert type(nba_nulls_) == type(nba_nulls)
pd.testing.assert_series_equal(nba_nulls, nba_nulls_)
### END HIDDEN TESTS
```

[Back to top](#Index:) 

## Problem 3

### Examining top 10 rows

**4 Points**

Use the `.head` method to display the first 10 rows of the nba data.  Assign your results as a DataFrame to `nba_head` below.

<!-- grade_id: cell-531911d9bad842be, points: 0, grade: false, solution: true -->
```python
### GRADED

nba_head = None

### BEGIN SOLUTION
nba_head = nba.head(10)
type(nba_head)
### END SOLUTION

# Answer check
print(type(nba_head))
print(nba_head.shape)
```

Output:
```
<class 'pandas.core.frame.DataFrame'>
(10, 9)
```

<!-- grade_id: cell-a1f9295805746d03, points: 4, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
nba_head_ = nba.head(10)
#
#
#
assert type(nba_head_) == type(nba_head)
pd.testing.assert_frame_equal(nba_head, nba_head_, 'Make sure to select the first 10 rows.')
### END HIDDEN TESTS
```